Social Experience Shapes Fighting Strategies in Drosophila

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

This work addresses an important question in the field of Drosophila aggression and mating. Prior social isolation is known to increase aggression in males, manifesting as increased lunging, which is suppressed by group housing (GH). However, it is also known that single housed (SH) males, despite their higher attempts to court females, are less successful. Here, Gao et al., develop a modified aggression assay to address this issue by recording aggression in Drosophila males for 2 hours, with a virgin female immobilized by burying its head in the food. They found that while SH males frequently lunge in this assay, GH males switch to higher intensity but very low frequency tussling. Constitutive neuronal silencing and activation experiments implicate cVA sensing Or67d neurons in promoting high frequency lunging, similar to earlier studies, whereas Or47b neurons promote low frequency but higher intensity tussling. Optogenetic activation revealed that three pairs of pC1SS2 neurons increase tussling. Cell-type-specific DsxM manipulations combined with morphological analysis of pC1SS2 neurons and side-by-side tussling quantification link the developmental role of DsxM to the functional output of these aggression-promoting cells. In contrast, although optogenetic activation of P1a neurons in the dark did not increase tussling, thermogenetic activation under visible light drove aggressive tussling. Using a further modified aggression assay, GH males exhibit increased tussling and maintain territorial control, which could contribute to a mating advantage over SH males, although direct measures of reproductive success are still needed.

Strengths:

Through a series of clever neurogenetic and behavioral approaches, the authors implicate specific subsets of ORNs and pC1 neurons in promoting distinct forms of aggressive behavior, particularly tussling. They have devised a refined territorial control paradigm, which appears more robust than earlier assays using a food cup (Chen et al., 2002). This new setup is relatively clutter-free and could be amenable to future automation using computer vision approaches. The updated Figure 5, which combines cell-type-specific developmental manipulation of pC1SS2 neurons with behavioral output, provides a link between developmental mechanisms and functional aggression circuits. The manuscript is generally well written, and the claims are largely supported by the data.

Thank you for the precise summary of the manuscript and acknowledgment of the novelty and significance of the study.

Weakness:

Although most concerns have been addressed, the manuscript still lacks a rigorous, objective method for quantifying lunging and tussling. Because scoring appears to have been done manually and a single lunge in a 30 fps video spans only 2-3 frames, the 0.2 s cutoff seems arbitrary, and there are no objective criteria distinguishing reciprocal lunging from tussling. Despite this, the study offers valuable insights into the neural and behavioral mechanisms of Drosophila aggression.

Thank you for this comment. The duration of each lunge was measured by analyzing the videos frame by frame—from the frame before the initiation of the lunge to the frame after its completion—resulting in an average span of 3–5 frames. Given a frame rate of 30 fps, this corresponds to approximately 0.1–0.17 seconds. We acknowledge that there are certain limitations for manually quantifying the two types of aggressive behaviors, which has now been stated in the newly added “Limitations of the Study” section in the revised manuscript.

Reviewer #2 (Public review):

Summary:

Gao et al. investigated the change of aggression strategies by the social experience and its biological significance by using Drosophila. Two modes of inter-male aggression in Drosophila are known: lunging, high-frequency but weak mode, and tussling, low-frequency but more vigorous mode. Previous studies have mainly focused on the lunging. In this paper, the authors developed a new behavioral experiment system for observing tussling behavior and found that tussling is enhanced by group rearing, while lunging is suppressed. They then searched for neurons involved in the generation of tussling. Although olfactory receptors named Or67d and Or65a have previously been reported to function in the control of lunging, the authors found that these neurons do not function in the execution of tussling and another olfactory receptor, Or47b, is required for tussling, as shown by the inhibition of neuronal activity and the gene knockdown experiments. Further optogenetic experiments identified a small number of central neurons pC1[SS2] that induce the tussling specifically. These neurons express doublesex (dsx), a sex-determination factor, and knockdown of dsx strongly suppresses the induction of tussling. In order to further explore the ecological significance of the aggression mode change in group-rearing, a new behavioral experiment was performed to examine the territorial control and the mating competition. And finally, the authors found that differences in the social experience (group vs. solitary rearing) and the associated change in aggression strategy are important in these biologically significant competitions. These results add a new perspective to the study of aggression behavior in Drosophila. Furthermore, this study proposes an interesting general model in which the social experience modified behavioral changes play a role in reproductive success.

Strengths:

A behavioral experiment system that allows stable observation of tussling, which could not be easily analyzed due to its low-frequency, would be very useful. The experimental setup itself is relatively simple, just the addition of a female to the platform, so it should be applicable to future research. The finding about the relationship between the social experience and the aggression mode change is quite novel. Although the intensity of aggression changes with the social experience was already reported in several papers (Liu et al., 2011 etc), the fact that the behavioral mode itself changes significantly has rarely been addressed, and is extremely interesting. The identification of sensory and central neurons required for the tussling makes appropriate use of the genetic tools and the results are clear. A major strength of this study in neurobiology is the finding that another group of neurons (Or47b-expressing olfactory neurons and pC1[SS2] neurons), distinct from the group of neurons previously thought to be involved in low-intensity aggression (i.e. lunging), function in the tussling behavior. Furthermore, the results showing that the regulation of aggression by pC1[SS2] neurons is based on the function of the dsx gene will bring a new perspective to the field. Further investigation of the detailed circuit analysis is expected to elucidate the neural substrate of the conflict between the two aggression modes. The experimental systems examining the territory control and the reproductive competition in Fig. 6 are novel and have advantages in exploring their biological significance. It is important to note that in addition to showing the effects of age and social experience on territorial and mating behaviors, the authors experimentally demonstrated that altered fighting strategy has effects with respect to these behaviors.

Thank you for your precise summary of our study and being very positive on the novelty and significance of the study.

Reviewer #3 (Public review):

In this revised manuscript, Gao et al. presented a series of well-controlled behavioral data showing that tussling, a form of high-intensity fighting among male fruit flies (Drosophila melanogaster) is enhanced specifically among socially experienced and relatively old males. Moreover, results of behavioral assays led authors to suggest that increased tussling among socially experienced males may increase mating success. They also concluded that tussling is controlled by a class of olfactory sensory neurons and sexually dimorphic central neurons that are distinct from pathways known to control lunges, a common male-type attack behavior.

A major strength of this work is that it is the first attempt to characterize behavioral function and neural circuit associated with Drosophila tussling. Many animal species use both low-intensity and high-intensity tactics to resolve conflicts. High-intensity tactics are mostly reserved for escalated fights, which are relatively rare. Because of this, tussling in the flies, like high-intensity fights in other animal species, have not been systematically investigated. Previous studies on fly aggressive behavior have often used socially isolated, relatively young flies within a short observation duration. Their discovery that 1) older (14-days old) flies tend to tussle more often than younger (2 to 7-days-old) flies, 2) group-reared flies tend to tussle more often than socially isolated flies, and 3) flies tend to tussle at later stage (mostly ~15 minutes after the onset of fighting), are the result of their creativity to look outside of conventional experimental settings. These new findings are key for quantitatively characterizing this interesting yet under-studied behavior.

Newly presented data have made several conclusions convincing. Detailed descriptions of methods to quantify behaviors help understand the basis of their claims by improving transparency. However, I remain concerned about authors' persistent attempt to link the high intensity aggression to reproductive success. The authors' effort to "tone down" the link between the two phenomena remains insufficient. There are purely correlational. I reiterate this issue because the overall value of the manuscript would not change with or without this claim.

Thank you for acknowledging the novelty and significance of the study. Regarding the relationship you mentioned between high-intensity aggression and reproductive success, we further toned down the statement between them throughout the manuscript in the revised manuscript. We also modified the title to “Social Experience Shapes Fighting Strategies in Drosophila”. In addition, we now added a ‘Limitations of the Study’ section to clearly state the correlation between tussling and reproductive success.

Reviewer #1 (Recommendations for the authors):

If possible, mention the EM-connectome data showing the minimal interneuronal path from Or47b ORNs to pC1SS2 neurons (even if derived from the female connectome), which can strengthen the model of parallel sensory-central pathways.

Thank you for this comment. According to data from the EM connectome, connecting Or47b ORNs to pC1d neurons requires at least two intermediate neurons. An example minimal pathway is: ORN_VA1v (L) → AL-AST1 (L) → PLP245 (L) → pC1d (R). We have added this point in the Discussion section of the revised manuscript.

I'm not convinced that labeling lunges as "gentle" combat behavior works, either in the abstract or elsewhere. While lunging is indeed a lower-intensity form of aggression compared to tussling, applying anthropomorphic descriptors risks misleading readers.

Thank you for this comment. We now use “low-intensity” instead of “gentle” to describe lunging.

In Materials & Methods, please cross-check all figure-panel references after the recent re-numbering (e.g. "Figure 5A6A" etc.).

Thank you for this comment. We have thoroughly verified the figure panel references in the Materials & Methods section.

Ensure that Table S1 is clearly cited in the main text where you first describe fly genotypes.

Thank you for this comment. We have now cited Table S1 in the main text.

There are multiple grammatical errors and typos throughout the manuscript. Please correct them. Some examples are below, but this is not an exhaustive list:

Line 98-102 requires rephrasing as the results are already published and not being observed by the authors.

Thank you for this comment. We have revised the manuscript to “we occasionally observed the high-intensity boxing and tussling behavior in male flies as previously reported (Chen et al., 2002; Nilsen et al., 2004), which….”

line 116- lower not 'lowed'.

Corrected.

line 942 & 945- knock-down males not 'knocking down males'.

Corrected. Thank you very much for these comments.

Reviewer #2 (Recommendations for the authors):

The authors have almost completely answered the major comments I have noted on the ver.1 manuscript: (1) They clearly show changes in fighting strategy in the territory control behavior experiment in Fig. 6-figure supplements. (2) A detailed description of how aggressive behavior is measured. Thus, I am convinced by this revision.

Thank you for these comments that make the manuscript a better version.

Furthermore, in Fig. 5, which examined the relationship of pC1[SS2] characteristics with the function of dsx, is a novel data and very interesting. I look forward to further developments.

Thank you. We will continue to explore this part in our future study.

However, one point still concerns me.

Line 192: Although the authors describe it as "usage-dependent," the trans-Tango technique is essentially a postsynaptic cell-labeling technique. It is possible that the labeling intensity in postsynaptic cells increases from the change in expression levels of the Or47b gene due to GH. However, there is no difference in the expression level of the Or47b gene labeled by GFP between SH and GH. Therefore, we cannot conclude that the expression of the Or47b gene is increased by rearing conditions.

The original paper on trans-TANGO (Talay et al., 2017) does not discuss the usage-dependency. A review of trans-synaptic labeling techniques (Ni, Front Neural Circuits. 2021) discusses that the increase in trans-TANGO signaling with aging may be related to synaptic strength, but there is no experimental evidence for this. In my opinion, the results in Figure 3-figure supplement 2 only weakly suggest that the increase in trans-TANGO signaling may be explained by an increase in synaptic strength due to group rearing.

We appreciate the reviewer’s insightful comment regarding the interpretation of the trans-Tango signal. Indeed, the original trans-Tango study (Talay et al., 2017) does not claim that the method is usage-dependent. The observed increase in trans-Tango labeling with age, as reported in their supplemental figures, may reflect accumulation over time, potentially influenced by synaptic maturation or increased component expression. To avoid overstating our results, we have revised the relevant statement in the manuscript to remove the term "usage-dependent" and now describe the change in trans-Tango signal more cautiously.

Reviewer #3 (Recommendations for the authors):

Below are the cases where their professed attempts to "tone down the statement" appear ignored:

Lines 27-29:

"Our findings... suggest how social experience shapes fighting strategies to optimize reproductive success".

We have now revised the manuscript to “Our findings… suggest that social experience may shape fighting strategies to optimize reproductive success.”

Lines 85-86:

"... discover that this infrequent yet intense form of combat is... crucial for territory dominance and mating competition".

We have now revised the manuscript to “…discover that this infrequent yet intense form of combat is enhanced by social enrichment, while the low-intensity lunging is suppressed by social enrichment.”

Lines 335-339:

"Here, we found that... GH males tend to... increase the high-intensity tussling, which enhances their territorial and mating competition."

We have removed “which enhances their territorial and mating competition” in the revised manuscript.

Lines 343-344:

"... presenting a paradox between social experience, aggression and reproductive success. Our result resolved this paradox..."

We have now revised the manuscript to “...Our results provide an explanation for this paradox…”

Lines 355-358:

"Interestingly, we found that the mating advantage gained through social enrichment can even offset the mating disadvantage associated with aging, further supporting the vital role of shifting fighting strategies in experienced, aged males."

We have removed “further supporting the vital role of shifting fighting strategies in experienced, aged males” in the revised manuscript.

Lines 361-362:

"These results separate the function of the two fighting forms and rectify out understanding of how social experiences regulate aggression and reproductive success."

We have removed this sentence in the revised manuscript.

Some may say that a speculative statement is harmless, but I think it indeed is harmful unless it is clearly indicated as a speculation. It is regrettable that authors remain reluctant to change their claim without providing any new supporting evidence. All three reviewers raised the same concern in the first round of review.

We apologize for not making the speculative nature of the statement clearer in the previous version. In the revised manuscript, we have now explicitly rephrased sentences to only suggest a correlation but not a causal link between tussling and reproductive success.

I have no choice but to keep my evaluation of the manuscript as "Incomplete" unless the authors thoroughly eliminate any attempt to link these two. This must go beyond changing a few words in the lines listed above.

Thank you for this comment. In addition to the lines listed above, we carefully checked all statements regarding the correlation between fighting strategies and reproductive success throughout the full text. Furthermore, we have also added a “Limitations of the Study” section to address the shortcomings of this study in the revised manuscript.

I do not have the same level of concern over the interpretation of Fig. 6A-C, because this is directly linked to aggressive interactions. Even if the socially isolated males do not engage in tussling, it is not a leap to assume that a different fighting tactic of socially experienced males can give them an advantage in defending a territory. To me, this is a sufficient ethological link with the observed behavioral change.

Thank you for this insightful comment.

The following are relatively minor, although important, concerns.

I beg to differ over the authors' definition of "tussling". Supplemental movies S1 and S2 appear to include "tussling" bouts in which 2 flies lunging at each other in rapid succession, and supplemental movie S3 appears to include bouts of "holding", in which one fly holds the opponent's wings and shakes vigorously. These cases suggest that the definition of "tussling" as opposed to "lunging" has a subjective element. However, I would not delve on this matter further because it is impossible to be completely objective over behavioral classification, even by using a computational method. An important point is that the definition is applied consistently within the publication. I have no reason to doubt that this was not the case.

Thank you for this comment. Since the analysis of tussling behavior was conducted manually, it is challenging to achieve complete objectivity. However, we made every effort to apply consistent criteria throughout the analysis. We have added a “Limitations of the Study” section in the revised manuscript to clearly state this caveat. We appreciate your understanding.

Authors now state that "all tester flies were loaded by cold anesthesia" (lines 432-433). I would like to draw attention to the well-known fact that anesthesia, whether by ice or by CO2, are long known to affect fly's subsequent behaviors (for aggression, see Trannoy S. et al., Learn. Mem. 2015. 22: 64-68). It will be prudent to acknowledge the possibility that this handling method could have contributed to unusually high levels of spontaneous tussling, which has not been reported elsewhere before.

Thank you for this comment. The increased tussling behavior observed in our study is unlikely due to cold anesthesia, as noted by Trannoy S. et al. (2015), cold anesthesia profoundly reduces locomotion and general aggressiveness in flies. We acknowledge that the use of cold anesthesia in behavioral experiments may have potential effects on aggression. To minimize this influence, we allowed the flies to recover and adapt for at least 30 minutes before behavioral recording. Moreover, both control and experimental groups were treated in exactly the same manner to ensure consistency.

It is intriguing that pC1SS2 neurons are dsx+ but fru-. Authors convincingly demonstrated that these neurons are clearly distinct from the P1a neurons, a well-characterized hub for male social behaviors. It is possible that pC1SS2 neurons overlap with previously characterized dsx+ neurons that are important for male aggressions (measured by lunges), such as in Koganezawa et al., Curr. Biol. 2016 and Chiu et al., Cell 2020, a point authors could have explicitly raised.

Thank you for this comment. We have added this point into the Discussion section of the revised manuscript, as follows: “That tussling-promoting… aggression (Koganezawa et al., 2016). Moreover, the anatomical features of pC1SS2 neurons are highly similar to the male-specific aggression-promoting (MAP) neurons identified by another previous study (Chiu et al., 2021).

I acknowledge the authors' courage to initiate an investigation to a less characterized, high intensity fighting behavior. Tussling requires the simultaneous engagement of two flies. Even if there are confusion over the distinction between lunges and tussling, authors' conclusion that socially experienced flies and socially isolated flies employ distinct fighting strategy is convincing. The concern I raised above is about the interpretation of the data, not about the quality of data.

Thank you for your constructive comments to make this manuscript better.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Social isolation generally increases aggression but decreases mating competition, resulting in an intricate and ambiguous relationship between social experience, aggression, and reproductive success. In male Drosophila, aggression is often characterized by lunging, a frequent and comparatively gentle combat behavior. Here, we provide a behavioral paradigm for studying a less frequent but more vigorous fighting form known as tussling. While social enrichment decreases lunging, aligning with past observations, it heightens the more forceful tussling behavior. These two forms of aggression rely on different olfactory receptor neurons, specifically Or67d for lunging and Or47b for tussling. We further identify three pairs of central pC1 neurons that specifically promote tussling. Moreover, shifting from lunging to tussling in socially enriched males is accompanied with better territory control and mating success. Our findings identify distinct sensory and central neurons for two fighting forms and suggest how social experience shapes fighting strategies to optimize reproductive success.

Revised version: Social isolation generally increases aggression but decreases mating competition, resulting in an intricate and ambiguous relationship between social experience, aggression, and reproductive success. In male Drosophila, aggression is often characterized by lunging, a frequent and comparatively low-intensity combat behavior. Here, we provide a behavioral paradigm for studying a less frequent but more vigorous fighting form known as tussling. While social enrichment decreases lunging, aligning with past observations, it heightens the more forceful tussling behavior. These two forms of aggression rely on different olfactory receptor neurons, specifically Or67d for lunging and Or47b for tussling. We further identify three pairs of central pC1 neurons that specifically promote tussling. Moreover, shifting from lunging to tussling in socially enriched males is accompanied with better territory control and mating success. Our findings identify distinct sensory and central neurons for two fighting forms and suggest that social experience may shape fighting strategies to optimize reproductive success.

Introduction

Competition for resources such as territory, food, and mates is crucial for animal survival and reproduction (Bergman et al., 2010; Hoffmann and Cacoyianni, 1990; Stockley and Campbell, 2013; Zwarts et al., 2012). Among diverse competitive strategies, physical combat is a prevalent form of fighting to repel rivals and regulated by both genetic factors and social environment (Chen, 2022; Zwarts et al., 2012). Understanding how social experience influences combat behaviors in animals to maximize their survival and reproduction remains a fundamentally important yet inadequately understood question.

In the fruit fly, Drosophila melanogaster, inter-male aggression includes relative low-intensity fights such as lunging, in which one fly rears up on its hind legs and snaps down onto its opponent, as well as relative high-intensity fights like tussling, where two males tumble over each other (Certel and Kravitz, 2012; Chen et al., 2002; Hoopfer, 2016; Palavicino-Maggio and Sengupta, 2022). A recent work showed that lunges are insufficient to establish dominance but mainly serves to maintain social status, while tussling frequence is correlated with dominance establishment (Simon and Heberlein, 2020). It is noteworthy that the majority of studies focus on the low-intensity lunging behavior, owing to its higher frequency and ease of observation. Substantial studies have identified regulatory genes and the neural circuit underlying the lunging-based aggressive behaviors (Hoopfer, 2016; Kravitz and Fernandez, 2015). In particular, the male-specific volatile pheromone 11-cis-vaccenyl acetate (cVA) acutely promotes lunging behavior via its olfactory receptor Or67d (Wang and Anderson, 2010). Interestingly, chronic cVA exposure in socially-enriched males reduces the level of lunging via another olfactory receptor Or65a (Liu et al., 2011). Consistently, it has been found that social isolation increases lunging frequency (Palavicino-Maggio and Sengupta, 2022; Yadav et al., 2024) by altering the expression of genes such as Cyp6a20 (Wang et al., 2008), Drosulfakinin (Dsk) (Agrawal et al., 2020), and Tachykinin (Tk) (Asahina et al., 2014). Social isolation also increases male aggression in mice (Matsumoto et al., 2005; Zelikowsky et al., 2018) and even in humans (Ferguson et al., 2005). In mice, the TK homologue Tachykinin 2 (Tac2) responds to social isolation and promotes inter-male aggressive behaviors (Zelikowsky et al., 2018). These findings contribute to the consensus that social isolation escalates aggression, whereas social enrichment mitigates it in animal models.

Meanwhile, social enrichment increases male mating success by enhancing the sensitivity of Or47b-expressing olfactory receptor neurons (ORNs) to fly pheromones common to both sexes (Lin et al., 2016; Sethi et al., 2019). Artificially activating Or47b ORNs also improves mating advantage in a competitive assay (Dweck et al., 2015). Group housing and age-dependent juvenile hormone (JH) jointly increase the expression of the male-specific proteins from fruitless (fruM), which is a master gene for male courtship behaviors (Peng et al., 2021; Sato et al., 2019; Yamamoto and Koganezawa, 2013), in the Or47b-expressing ORNs (Hueston et al., 2016; Sethi et al., 2019; Zhao et al., 2020). Notably, courtship levels in single-housed (SH) wild-type males are equal to (Dankert et al., 2009; Pan and Baker, 2014), in some cases, even stronger than those in group-housed (GH) males (Dankert et al., 2009). In fact, group-housing decreases excitability of the courtship-promoting P1 neurons (Inagaki et al., 2014), which express both terminal genes in the sex determination hierarchy, fru and doublesex (dsx). These findings raise a question on how social enrichment, which suppresses lunging-based aggression and does not increase courtship intensity, would facilitate mating competition.

Earlier version: Here, we modify a behavioral paradigm for studying inter-male tussling in Drosophila and discover that this infrequent yet intense form of combat is enhanced by social enrichment and is crucial for territory dominance and mating competition. This process requires olfactory receptor Or47b, but not Or67d or Or65a. We further identify three pairs of central dsx-expressing pC1 neurons that specifically promote tussling. Moreover, we demonstrate that socially enriched males can overcome mating disadvantages associated with aging.

Revised version: Here, we provide a behavioral paradigm for studying inter-male tussling in Drosophila and discover that this infrequent yet intense form of combat is enhanced by social enrichment, while the low-intensity lunging is suppressed by social enrichment. This social enhancement of tussling requires olfactory receptor Or47b, but not Or67d or Or65a. We further identify three pairs of central dsx-expressing pC1 neurons that specifically promote tussling. Moreover, we demonstrate that socially enriched males can overcome mating disadvantages associated with aging.

Results

A Paradigm for High-intensity Aggression in Drosophila

Earlier version: While studying the lunging-based aggression in Drosophila using a high-throughput aggression assay based on previously aggression paradigms (Certel and Kravitz, 2012; Wu et al., 2020; Zhou et al., 2008), we occasionally observed the high-intensity boxing and tussling behavior in male flies (Chen et al., 2002; Nilsen et al., 2004), which are difficult to distinguish and hereafter simply referred to as tussling. We reasoned that the intensity, instead of the frequency, could be decisive for a combat (Simon and Heberlein, 2020). Thus, although tussling is a less frequent behavior, it could serve a vital function.

Revised version: While studying the lunging-based aggression in Drosophila using a high-throughput aggression assay based on previously aggression paradigms (Certel and Kravitz, 2012; Wu et al., 2020; Zhou et al., 2008), we occasionally observed the high-intensity boxing and tussling behavior in male flies as previously reported (Chen et al., 2002; Nilsen et al., 2004), which are difficult to distinguish and hereafter simply referred to as tussling. We reasoned that the intensity, instead of the frequency, could be decisive for a combat (Simon and Heberlein, 2020). Thus, although tussling is a less frequent behavior, it could serve a vital function.

Earlier version: To induce a higher level of tussling for further manipulation, we first introduced a fixed live virgin female in the same behavioral chamber for assaying lunging. The female fly was immobilized in the middle of the food, with its head embedded in the food and its body exposed to enhance male competition (Figure 1A). We found that males that were group-housed for 7 days showed a significantly higher frequency of tussling behavior and reduced latency to tussle in the presence of an embedded virgin female (Figure 1B). We next compared tussling in male flies of different ages. We observed a significantly higher proportion of 14-day-old males raised in groups that showed tussling behavior during the 2-hour observation period, compared to the 7-day-old males (Figure 1B). Furthermore, 2-day-old males never displayed tussling behavior. In contrast, 14-day-old males showed lowed level of lunging behavior than 7-day-old males did (Figure 1-figure supplement 1). Previous studies have shown that, in addition to mating resources, food quality is another crucial factor affecting aggressive behaviors (Lim et al., 2014). Therefore, we next assayed tussling behavior in 14-day-old males in the presence of an embedded virgin female with different concentrations of sucrose or yeast in the food during the behavioral test. We found that these males displayed similar levels of tussling behavior with different concentrations of sucrose in the food (Figure 1C). On the contrary, male tussling intensity increased significantly with higher yeast concentrations such that more than 70% of males displayed this high-intensity fighting behavior (Figures 1D, see Figure 1-video supplement 1). Taken together, these results provide an efficient behavioral paradigm for the high-intensity tussling behavior and indicate that older males display more tussling but less lunging than younger ones.

Revised version: To induce a higher level of tussling for further manipulation, we first introduced a fixed live virgin female in the same behavioral chamber for assaying lunging. The female fly was immobilized in the middle of the food, with its head embedded in the food and its body exposed to enhance male competition (Figure 1A). We found that males that were group-housed for 7 days showed a significantly higher frequency of tussling behavior and reduced latency to tussle in the presence of an embedded virgin female (Figure 1B). We next compared tussling in male flies of different ages. We observed a significantly higher proportion of 14-day-old males raised in groups that showed tussling behavior during the 2-hour observation period, compared to the 7-day-old males (Figure 1B). Furthermore, 2-day-old males never displayed tussling behavior. In contrast, 14-day-old males showed a lower level of lunging behavior than 7-day-old males did (Figure 1-figure supplement 1). Previous studies have shown that, in addition to mating resources, food quality is another crucial factor affecting aggressive behaviors (Lim et al., 2014). Therefore, we next assayed tussling behavior in 14-day-old males in the presence of an embedded virgin female with different concentrations of sucrose or yeast in the food during the behavioral test. We found that these males displayed similar levels of tussling behavior with different concentrations of sucrose in the food (Figure 1C). On the contrary, male tussling intensity increased significantly with higher yeast concentrations such that more than 70% of males displayed this high-intensity fighting behavior (Figures 1D, see Figure 1-video supplement 1). Taken together, these results provide an efficient behavioral paradigm for the high-intensity tussling behavior and indicate that older males display more tussling but less lunging than younger ones.

A paradigm for tussling behavior in male Drosophila.

(A) Diagrams of two fighting paradigms for Drosophila tussling behavior. The left diagram shows a traditional fighting paradigm with no female (NF) in the food area and a modified fighting paradigm with a fixed female (FF) in the center of the food area. Before the experiment began, two male flies were placed on the upper layer, separated from the food by a film. After a 30-minute adaptation, the film was removed to start recording. The right diagram shows a pair of tussling male Drosophila. All tester males in this figure were group housed.

(B) The influence of age and female factors on male tussling behavior. From top to bottom, the tussle latency, the number of tussles, and the proportion of tussles occurring are shown respectively. n=30, 30, 30, 46, 46, 46 from left to right. Bars sharing the same letter are not significantly different according to Kruskal-Wallis test (tussle latency, number of tussles) and Chi-square test (proportion of tussles occurring).

(C) Sucrose concentration of food have no effect on male tussling behavior. n=30 for each group, n.s., not significant, Kruskal-Wallis test (tussle latency, number of tussles), and Chi-square test (proportion of tussles occurring).

(D) The influence of yeast concentration on male tussling behavior. n=30 for each group, n.s., not significant, *P < 0.05, ***P < 0.001, ****P < 0.0001, Kruskal-Wallis test (tussle latency, number of tussles), and Chi-square test (proportion of tussles occurring). Error bars correspond to SEM.

Social Enrichment Inhibits Low-intensity Lunging but Enhances High-intensity Tussling

Previous studies have shown that social enrichment inhibits aggressive behaviors, specifically the low-intensity lunging, in male flies (Yadav et al., 2024). Therefore, we wondered whether social enrichment might also inhibit the high-intensity tussling behavior in males. To answer this question, we compared both fighting behaviors in SH and GH 14-day-old males (Figure 2A). Lunging and tussling events from each group of representative samples were exhibited in raster plots for comparison (Figure 2B). We scored lunging behavior mainly in the first 10-min of test, due to its relatively higher occurrence, and tussling behavior during the whole 2-hour observational period since males take much longer time to initiate tussling (Figure 1D). Interestingly, we found that while SH males performed more frequent lunging behavior in 10 min than GH males did, which is consistent with previous findings, they rarely showed tussling behavior. In contrast, GH males displayed a large number of tussling events within 2 hours (Figure 2B). Statistical analysis of tussling and lunging behaviors showed that GH males exhibited reduced tussling latency and enhanced tussling frequency, in addition to increased lunging latency and decreased lunging frequency, compared to SH males (Figures 2C-F). We also quantified the duration of each tussling and lunging events, and found that the average duration of individual lunging events was less than 0.2 second, while a tussling event lasted from seconds to minutes (Figure 2G). This further suggests that tussling is a more intense form of aggressive behavior compared to lunging. These results indicate that social enrichment inhibits low-intensity lunging but promotes high-intensity tussling in male flies.

Social enrichment inhibits lunging behavior but promotes tussling behavior in males.

(A) Schematic diagram of Drosophila rearing conditions and experimental procedure.

(B) Representative raster plots illustrating lunge events in the first 10 minutes and tussle events over 2 hours in single-housed (SH) and group-housed (GH) males.

Removed in the revised version: (C, D) Latency (C) and number (D) of tussling behavior in SH and GH males, n=27, 30 for SH and GH males, respectively, ****P < 0.0001, Mann-Whitney test for both.

Earlier version: (E, F) Latency (E) and number (F) of lunging behavior in SH and GH males, n=27, 30 for SH and GH males, respectively, ****P < 0.0001, Mann-Whitney test for both. Error bars correspond to SEM.

Revised version: (C, D) Latency (C) and number (D) of tussling behavior in SH and GH males, n=27, 30 for SH and GH males, respectively, ****P < 0.0001, Mann-Whitney test for both. (E, F) Latency (E) and number (F) of lunging behavior in SH and GH males, n=27, 30 for SH and GH males, respectively, ****P < 0.0001, Mann-Whitney test for both. Error bars correspond to SEM.

(G) Average duration of lunging events in the first 10 minutes and tussling events within 2 hours in GH males, n=30 events from 10 samples for each group, ****P < 0.0001, Mann-Whitney test. Error bars correspond to SEM.

Or47b ORNs Mediate Social Experience-induced Tussling Behavior

Earlier version: To further investigate the sensory basis underlying the opposite regulation of lunging and tussling behaviors by social experience, we tried to examine the role of cVA-sensitive neurons, as well as other fruM-expressing ORNs, in mediating intermale aggression. Two cVA sensory neurons, Or67d and Or65a ORNs, play crucial roles in enhancing the lunging-based aggression in SH males, in which Or67d ORNs mediate acute cVA stimulation and promote male aggression (Wang and Anderson, 2010), while Or65a ORNs mediate chronic cVA stimulation and inhibit male aggression in GH males (Liu et al., 2011). In addition, the fruM-expressing Or47b ORNs are sensitive to a class of pheromones common to both sexes, such as palmitoleic acid (PA) (Lin et al., 2016), and their sensitivity is enhanced under group rearing conditions (Sethi et al., 2019). Ir84a ORNs, which are also fruM-positive, are involved in the perception of mixed food odors and promote male courtship behavior (Grosjean et al., 2011). We silenced Or67d, Or65a, Or47b or Ir84a ORNs by expressing the inwardly rectifying potassium channel Kir2.1 (Baines et al., 2001) and assayed lunging and tussling behaviors in 14-day-old GH males. We found that silencing Or47b ORNs, but not Or67d, Or65a or Ir84a ORNs, significantly decreased male tussling, compared to the control (Figure 3A, B). In contrast, silencing Or67d ORNs, but not other ORNs, significantly reduced male lunging (Figure 3C, D). These loss-of-function results suggest that Or47b ORNs mediate male tussling behavior, whereas Or67d ORNs regulate male lunging behavior.

Revised version: To further investigate the sensory basis underlying the opposite regulation of lunging and tussling behaviors by social experience, we tried to examine the role of cVA-sensitive neurons, as well as other fruM-expressing ORNs, in mediating intermale aggression. Two cVA sensory neurons, Or67d and Or65a ORNs, play crucial roles in enhancing the lunging-based aggression in SH males, in which Or67d ORNs mediate acute cVA stimulation and promote male aggression (Wang and Anderson, 2010), while Or65a ORNs mediate chronic cVA stimulation and inhibit male aggression in GH males (Liu et al., 2011). In addition, the fruM-expressing Or47b ORNs are sensitive to a class of pheromones common to both sexes, such as palmitoleic acid (PA) (Lin et al., 2016), and their sensitivity is enhanced under group rearing conditions (Sethi et al., 2019). Ir84a ORNs, which are also fruM-positive, are involved in the perception of mixed food odors and promote male courtship behavior (Grosjean et al., 2011). We silenced Or67d, Or65a, Or47b or Ir84a ORNs by expressing the inwardly rectifying potassium channel Kir2.1 (Baines et al., 2001) and assayed lunging and tussling behaviors in 14-day-old GH males. We found that silencing Or47b ORNs, but not Or67d, Or65a or Ir84a ORNs, significantly decreased male tussling, compared to the control (Figure 3A, B, see Table S1 for detailed genotypes). In contrast, silencing Or67d ORNs, but not other ORNs, significantly reduced male lunging (Figure 3C, D). These loss-of-function results suggest that Or47b ORNs mediate male tussling behavior, whereas Or67d ORNs regulate male lunging behavior. To elucidate the potential role of the Or47b receptor in mediating tussling behavior, we knocked down its expression in Or47b ORNs using RNA interference (RNAi). Our results showed a significant decrease in male tussling behavior when the Or47b receptor expression was reduced (Figure 3E, F), indicating that the functionality of Or47b ORNs in tussling is dependent on the Or47b receptor. As previous studies suggested that the male-specific fruM sustains Or47b expression in an activity-dependent manner (Sethi et al., 2019), we next expressed microRNAs targeting fruM (UAS-fruMi) (Meissner et al., 2016) in Or47b ORNs and assayed male tussling. The efficiency of Or47b and fruM RNAi knockdown was both validated by quantitative PCR (Figure 3-figure supplement 1). We observed a slight increase of tussling latency in fruM knocked-down males, although not significantly different from the control (Figure 3E, F), suggesting that Or47b, rather than FruM, plays a more crucial role in mediating male tussling.

Or47b neurons are required for male tussling but not lunging.

(A, B) The effect of inactivating different sensory neurons on latency (A) and number

(B) of tussling behavior in 14-day-old group-housed males (G14). n=30 for each group, n.s., not significant, **P < 0.01, Kruskal-Wallis test.

(C, D) The effect of inactivating different sensory neurons on the latency (C) and number (D) of lunging behavior in G14 males. n=30 for each group, n.s., not significant, **P < 0.01, ***P < 0.001, Kruskal-Wallis test.

(E, F) The effect of knocking down Or47b or fruM expression in Or47b ORNs on the latency (E) and number (F) of tussling behavior in G14 males, respectively. n=30 for each group, n.s., not significant, ****P < 0.0001, Kruskal-Wallis test.

(G, H) The effect of activating Or47b ORNs on latency (G) and number (H) of tussling behavior under conditions of GH or SH for 14 days, n=60, 60, 60, 49 respectively, n.s., not significant, ****P < 0.0001, Kruskal-Wallis test. Error bars correspond to SEM.

Removed in the revised version: To elucidate the potential role of the Or47b receptor in mediating tussling behavior, we knocked down its expression in Or47b ORNs using RNA interference (RNAi). Our results showed a significant decrease in male tussling behavior when the Or47b receptor expression was reduced (Figure 3E, F), indicating that the functionality of Or47b ORNs in tussling is dependent on the Or47b receptor. As previous studies suggested that the male-specific fruM sustains Or47b expression in an activity-dependent manner (Sethi et al., 2019), we next expressed microRNAs targeting fruM (UAS-fruMi) (Meissner et al., 2016) in Or47b ORNs and assayed male tussling. The efficiency of Or47b and fruMRNAi knockdown was both validated by quantitative PCR (Figure 3-figure supplement 1). We observed a slight increase of tussling latency in fruM knocked-down males, although not significantly different from the control (Figure 3E, F), suggesting that Or47b, rather than FruM, plays a more crucial role in mediating male tussling.

Earlier version: To further confirm that the effect of social enrichment on tussling depends on the activity of Or47b ORNs, we artificially activated Or47b ORNs by expressing the bacterially derived sodium channel (NaChBac) (Ren et al., 2001) in 14-day-old SH males. We found that these SH males with enhanced activity of Or47b ORNs showed a significant increase in tussling behavior compared to control SH males. Interestingly, the level of tussling in these SH males with activated Or47b ORNs was indistinguishable from that in GH males (Figure 3G, H). These gain-of-function results indicate that social enrichment acts through Or47b ORNs to promote male tussling. Using the trans-Tango technique that labels downstream neurons in a usage-dependent manner (Talay et al., 2017), we found substantially more downstream signals of Or47b ORNs in GH males than those in SH males (Figure 3-figure supplement 2), further supporting the involvement of Or47b ORNs under social enrichment.

Revised version: To further confirm that the effect of social enrichment on tussling depends on the activity of Or47b ORNs, we artificially activated Or47b ORNs by expressing the bacterially derived sodium channel (NaChBac) (Ren et al., 2001) in 14-day-old SH males. We found that these SH males with enhanced activity of Or47b ORNs showed a significant increase in tussling behavior compared to control SH males. Interestingly, the level of tussling in these SH males with activated Or47b ORNs was indistinguishable from that in GH males (Figure 3G, H). These gain-of-function results indicate that social enrichment acts through Or47b ORNs to promote male tussling. Using the trans-Tango technique that labels downstream neurons through ligand-receptor mediated signaling (Talay et al., 2017), we found substantially more downstream signals of Or47b ORNs in GH males than those in SH males (Figure 3-figure supplement 2), suggesting the involvement of Or47b ORNs under social enrichment.

Distinct Central Brain Neurons for Lunging and Tussling

We next set out to identify the central neurons involved in regulating the high-intensity tussling. We found that silencing the mushroom body neurons by expressing an inwardly rectifying potassium channel Kir2.1 (Baines et al., 2001) driven by R19B03-GAL4 did not affect male tussling, suggesting that the experience-dependent tussling behavior is not dependent on the classical learning and memory center. Surprisingly, inactivation of the previously-identified aggression-promoting P1a and TK neurons did not affect tussling behavior (Figure 4A, B). These results imply that the central circuit for tussling behavior may be distinct from the previously reported aggression circuits. As dsx-expressing pC1 neurons, also termed P1 if co-expressing fruM, are crucial for sexual and aggressive behaviors in both males and females, we also silenced two subsets of pC1 neurons (pC1SS1 and pC1SS2), which were sexually dimorphic, and their functions were previously characterized in females but not in males (Wang et al., 2021). We found that inactivation of pC1SS2 but not pC1SS1 neurons significantly reduced the tussling behavior in male flies (Figure 4A, B). The cell bodies of pC1SS2 and P1a neurons are both located in the center of posterior brain and project to overlapping but distinct parts of the lateral protocerebral complex (Figure 4C-E). These results suggest that pC1SS2 neurons may be key central neurons regulating tussling behavior.

A subset of pC1 neurons specifically promote male tussling behavior.

(A, B) The effect of inactivating candidate interneurons on latency (A) and number (B) of male tussling behavior, n=30 for each group, bars sharing the same letter are not significantly different according to Kruskal-Wallis test.

(C-E) The expression pattern of P1a-spGAL4 (C) and pC1SS2-spGAL4 (D) in male brain, and (E) the registration of P1a (magenta) and pC1SS2 (green) neurons in a standard brain. The GFP expression indicates P1a (magenta) and pC1SS2 (green) neurons in the male brain counterstained with nc82 (blue). White dashed boxes in (C-E) were zoomed in and shown below each panel. Scale bars, 100μm.

(F-H) Statistics on tussling (F), lunging (G) and unilateral wing extension (UWE, H) by UAS-CsChrimson/+ males. Red light activation with a light intensity of 0.02 mW/mm2 for 1 minute was applied at the 2nd and 5th minutes of this and subsequent experiments. n=32 for each group, n.s., not significant, Friedman test.

(I-K) Statistics on tussling (I), lunging (J) and UWE (K) in P1a>UAS-CsChrimson males. n=32 for each group, n.s., not significant, ****P < 0.0001, Friedman test.

(L-N) Statistics on tussling (L), lunging (M) and UWE (N) in pC1SS2>UAS-CsChrimson males. n=32 for each group, n.s., not significant, ***P < 0.001, ****P < 0.0001, Friedman test.

To further compare the function of P1a and pC1SS2 neurons in regulating male behaviors, we performed optogenetic activation experiments on these neurons expressing CsChrimson. Compared to control male pairs that did not show any significant change upon red light stimulation (Figure 4F-H), optogenetic activation of P1a neurons induced acute courtship behavior as indicated by unilateral wing extension (UWE), and upon red light removal, there was a significant increase in lunging but not tussling (Figure 4I-K), consistent with previous findings. In contrast, optogenetic activation of pC1SS2 neurons induced strong and acute tussling behavior, with more than 50% of the time during red light in tussling, and the tussling behavior ceased immediately upon removal of red light (Figure 4L, see Figure 4-video supplement 1). Activation of pC1SS2 neurons did not significantly affect either male lunging or courtship (Figure 4M, N). These results indicate that pC1SS2 neurons acutely promote the high-intensity tussling behavior.

Previous studies have revealed a crucial role of P1a neurons in promoting a persistent internal state that enhances male aggression (Hoopfer et al., 2015). To test whether P1a neurons can promote both lunging and tussling behaviors, we further performed thermal activation experiments. We found that activation of P1a neurons via dTrpA1 at 30°C induced male tussling, just like activation of pC1SS2 neurons did (Figure 4-figure supplement 1). As thermal activation but not optogenetic activation of P1a neurons induced tussling, we reasoned that vision might be a key factor for this discrepancy, since all optogenetic experiments were performed in dark or red light. Thus, we re-performed these thermal activation experiments in dark and found that activation of pC1SS2 but not P1a neurons induced tussling, although the level of tussling by pC1SS2-activated males was lower compared to that in light, suggesting a crucial role of vision in male tussling. These optogenetic and themogenetic activation experiments suggest that P1a neurons can promote both lunging and tussling behaviors, but their function on tussling requires visual stimulation, while pC1SS2 neurons play a more direct role in acutely promoting male tussling.

We next used retro-Tango (Sorkaç et al., 2023) and trans-Tango (Talay et al., 2017) techniques to explore the potential upstream and downstream circuits of P1a and pC1SS2 neurons, and observed substential differences in both conditions (Figure 4-figure supplement 2), suggesting largely parallel inputs and outputs for P1a and pC1SS2 neurons.

DsxM Regulates the Development and Function of pC1SS2 Neurons

To further compare the differences between P1a neurons and pC1SS2 neurons, we investigated whether P1a and pC1SS2 neurons express fru and/or dsx using FruM and DsxM antibodies (Chen et al., 2021; Peng et al., 2022). We found that most P1a neurons were fru-positive and dsx-positive, while pC1SS2 neurons were dsx-positive but fru-negative (Figure 5A-5D). These results suggest that dsx instead of fru may play crucial roles in the development and function of pC1SS2 neurons.

DsxM regulates the development and tussling-promoting function of pC1SS2 neurons.

(A, B) The co-staining of P1a or pC1SS2 neurons with the FruM antibody (A), and the quantification of fru-positive neurons in P1a and pC1SS2 populations (B). n=5 for each group. Scale bars, 50μm.

(C, D) The co-staining of P1a or pC1SS2 neurons with the DsxM antibody, and The quantification of dsx-positive neurons in P1a and pC1SS2 populations (D), n=9 for each group. Scale bars, 50μm.

Earlier version: (E-K) Knocking down dsx in pC1SS2 neurons impact their development and tussling-promoting function. (E, F) Immunostaining of pC1SS2 neurons in control (E) and dsx knocking-down (F) males. Magenta represents nc82 signals, and green represents GFP signals. The images within the yellow dashed box and white dashed box are magnified 2× and 4×, respectively, and displayed on the right. Scale bars, 100μm. (G) Quantification of the number of pC1SS2 neurons in control and dsx knocking-down males. n=10, 9 for control and experimental group respectively, *P < 0.05, Mann-Whitney test. Error bars correspond to SEM. (H) Quantification of the neuronal projection length shown in the 2× magnified images in (E) and (F). n=10 for each group, ****P < 0.0001, Mann-Whitney test. (I-K) Tussling time by pC1SS2> UAS-CsChrimson (I) and pC1SS2 > UAS-CsChrimson & UAS-dsxM-RNAi (J) males, which were quantified in (K). Red light activation with a light intensity of 0.02 mW/mm2 for 1 minute was applied at the 2nd and 5th minutes. n=12, bars sharing the same letter are not significantly different. Friedman test for (I) and (J), Kruskal-Wallis test for (K).

Revised version: (E-K) Knocking down dsx in pC1SS2 neurons impact their development and tussling-promoting function. (E, F) Immunostaining of pC1SS2 neurons in control (E) and dsx knock-down (F) males. Magenta represents nc82 signals, and green represents GFP signals. The images within the yellow dashed box and white dashed box are magnified 2× and 4×, respectively, and displayed on the right. Scale bars, 100μm. (G) Quantification of the number of pC1SS2 neurons in control and dsx knock-down males. n=10, 9 for control and experimental group respectively, *P < 0.05, Mann-Whitney test. Error bars correspond to SEM. (H) Quantification of the neuronal projection length shown in the 2× magnified images in (E) and (F). n=10 for each group, ****P < 0.0001, Mann-Whitney test. (I-K) Tussling time by pC1SS2 > UAS-CsChrimson (I) and pC1SS2 > UAS-CsChrimson & UAS-dsxM-RNAi (J) males, which were quantified in (K). Red light activation with a light intensity of 0.02 mW/mm2 for 1 minute was applied at the 2nd and 5th minutes. n=12, bars sharing the same letter are not significantly different. Friedman test for (I) and (J), Kruskal-Wallis test for (K).

To investigate the function of dsx in pC1SS2 neurons, we knocked down dsx via RNAi and compared the morphological and functional changes of these neurons. We found that dsxM knockdown in pC1SS2 neurons slightly increased the number of cells from 2-3 pairs to 3-4 pairs (Figures 5E-5G), and significantly lengthened the axonal projections ventrally (Figure 5H). We reasoned that these developmental abnormality may effect the tussling-promoting functions of pC1SS2 neurons, and tried to activate these neurons with or without knocking down dsx. We found that males with wild-type dsx function displayed intensive tussling behavior for approximately 50% and 70% of the total activation time, respectively, during the first and second activation periods, with the response returning to baseline levels after the removal of red light (Figure 5I, 5K). However, males with dsx knocked down in pC1SS2 neurons showed no tussling behavior during the first activation session, and few tussling behavior during the second activation session (Figure 5J, 5K). Together, these results reveal the crucial role of the dsx gene in the development and tussling-promoting function of pC1SS2 neurons.

Territory and Mating Advantage in Tussling-favored Experienced Males

Earlier version: As our results showed that GH males and SH males were more likely to perform the low-frequency, high-intensity tussling and the high-frequency, low intensity lunging, respectively, we wondered which fighting strategy was more practical in the competition for territory and mating resources.

Revised version: As our results showed that GH males and SH males were more likely to perform the low-frequency, high-intensity tussling and the high-frequency, low intensity lunging, respectively, we wondered which fighting strategy might be more practical in the competition for territory and mating resources.

Earlier version: To address this question, we first designed an experimental paradigm for territory competition (Figure 6A). A circular territory, which is provided with food and a fixed virgin female and separated by a wall with a small door open to the outer area, was competed by two tester males, one of which was marked on their wings. We manually analyzed the first 10 winning events, each was defined as one male successfully displacing the other from the inner resource area. The winning index was calculated as the proportional difference in the number of wins between the marked male and unmarked male. In control groups such as two GH males of the same age, we observed some ‘winner takes all’ phenomenon (one male won nearly all of the 10 encounters), but overall, the winning index was close to zero due to the randomness of marking (see methods). We further found that while GH 7-day-old males (G7) won slightly, but not significantly, more than SH 7-day-old males (S7) (Figure 6B), older G14 males had a significant advantage over the S14 males (Figure 6C and Figure 6-figure supplement 1, Table S2, see Figure 6-video supplement 1). Additionally, we observed that when G14 males were paired with S14 males, the proportion of males performing tussling behavior was not significantly different from that between two G14 males. However, the duration of each tussling event in G14 and S14 male pairs was significantly shorter than that between two G14 males (Figure 6-figure supplement 2).

Revised version: To at least partially address this question, we first designed an experimental paradigm for territory competition (Figure 6A). A circular territory, which is provided with food and a fixed virgin female and separated by a wall with a small door open to the outer area, was competed by two tester males, one of which was marked on their wings. We manually analyzed the first 10 winning events, each was defined as one male successfully displacing the other from the inner resource area. The winning index was calculated as the proportional difference in the number of wins between the marked male and unmarked male. In control groups such as two GH males of the same age, we observed some ‘winner takes all’ phenomenon (one male won nearly all of the 10 encounters), but overall, the winning index was close to zero due to the randomness of marking (see methods). We further found that while GH 7-day-old males (G7) won slightly, but not significantly, more than SH 7-day-old males (S7) (Figure 6B), older G14 males had a significant advantage over the S14 males (Figure 6C and Figure 6-figure supplement 1, Table S2, see Figure 6-video supplement 1). Additionally, we observed that when G14 males were paired with S14 males, the proportion of males performing tussling behavior was not significantly different from that between two G14 males. However, the duration of each tussling event in G14 and S14 male pairs was significantly shorter than that between two G14 males (Figure 6-figure supplement 2).

Social enrichment enhances male territorial control and mating competition.

(A) Schematic diagram of the paradigm for testing territorial control in males.

(B) The territorial control test of 7-day-old GH (G7) and SH (S7) males. G7a and G7b stand for marked G7 male and unmarked G7 male, respectively. n=12 for each group, n.s., not significant, one sample t test.

(C) The territorial control test of 14-day-old GH (G14) and SH (S14) males. n=12 for each group, n.s., not significant, *P < 0.05, one sample t test.

(D) Schematic diagram of the paradigm for testing male mating advantage.

(E) The mating competition test between G7 and S7 males. n=12 for each group, n.s., not significant, one sample t test.

(F) The mating competition test between G14 and S14 males. n=12 for each group, n.s., not significant, *P < 0.05, one sample t test.

(G and H) The mating competition test between young males and old males, n=12 for each group, n.s., not significant, *P < 0.05, ****P < 0.0001, one sample t test for comparison within group (red labeling), Mann-Whitney test for comparison between groups (black labeling). Error bars correspond to SEM.

Earlier version: Next, we conducted further investigations into the role of social experience and the two types of fighting behaviors in mating competition (Figure 6D). Two males and one virgin female were loaded into the behavioral chamber at the same time and recorded for courtship behaviors. A winner was defined as the male that first copulated with the female. We used the copulation advance index, which was calculated as the proportional difference between marked males and unmarked males in winning copulation (see methods). We found that G7 males had no significant advantage over S7 males (Figure 6E); however, G14 males displayed a significantly increased mating advantage over S14 males (Figure 6F, see Figure 6-video supplement 2). Together, these results suggest that the tussling-favored GH males are more competitive in territory control and mating compared to the lunging-favored SH males.

Revised version: Next, we conducted further investigations into the role of social experience and the two types of fighting behaviors in mating competition (Figure 6D). Two males and one virgin female were loaded into the behavioral chamber at the same time and recorded for courtship behaviors. A winner was defined as the male that first copulated with the female. We used the copulation advance index, which was calculated as the proportional difference between marked males and unmarked males in winning copulation (see methods). We found that G7 males had no significant advantage over S7 males (Figure 6E); however, G14 males displayed a significantly increased mating advantage over S14 males (Figure 6F, see Figure 6-video supplement 2). Together, these results suggest that the tussling-favored GH males are more competitive in territory control and mating compared to the lunging-favored SH males, supporting a correlation between tussling and reproductive competetion.

The above results showed that GH males were more competitive compare to SH males, especially in older males, suggesting that social experience may play a potential role in altering age-related mating competition. To test this hypothesis, we first tested mating competition between males of different ages under the same rearing conditions. The results showed that younger males, whether reared in group or isolation, had a greater mating advantage compared to older males (G7 > G14, S7 > S14, G14 > G21, S14 > S21) (Figure 6G, Table S2). These results indicate that there is generally a disadvantage in mating competition associated with aging. We next investigated whether the mating advantage induced by social experience could compensate for the mating disadvantage caused by aging. We tested pairs of males with different ages and rearing conditions (G7 vs. S14, S7 vs. G14, G14 vs. S21, S14 vs. G21) for mating competition. We found that while 14-day-old males were less competitive than 7-day-old males regardless of the rearing condition, indicative of the aging-related disadvantage, 21-day-old GH males were significantly more competitive than 14-day-old SH males (G21 > S14) (Figure 6H).

To further investigate whether the territorial control and mating competition advantage in GH males are related to their high tussling behavior, we paired males with low tussling behavior (Or47b-GAL4/UAS-Kir2.1 and pC1SS2/UAS-Kir2.1) with control males for territorial control and mating competition tests. We found that males with inactivated Or47b or pC1SS2 neurons were incompetitive for both territorial control and mating competition compared to control males (Figure 6-figure supplement 3), further supporting a correlation between the occurrence of tussling behavior and the acquisition of territorial and mating competition advantages.

Taken together, these results indicate that long-term social enrichment in older males increases their territorial and mating competition, which can overcome their mating disadvantage associated with aging (Figure 7).

Summary model of fighting strategies and reproductive success in male Drosophila.

Earlier version: Age and social enrichment inhibit the low-intensity, high-frequency lunging while promoting the low-frequency, high-intensity tussling. This shift in fighting strategies among experienced, aged males can enhance their territorial control and mating competition, even offsetting aging-related mating disadvantages. Lunging and tussling behaviors are regulated by distinct sensory and central neurons.

Revised version: Age and social enrichment inhibit the low-intensity, high-frequency lunging while promoting the low-frequency, high-intensity tussling. This shift in fighting strategies among experienced, aged males is accompanished with better territorial control and mating competition, even offsetting aging-related mating disadvantages. Lunging and tussling behaviors are regulated by distinct sensory and central neurons.

Discussion

Earlier version: Regulation of physiology and behavior by social experience is an important feature across animal species. Here, we found that social experience shapes the fighting strategies of male flies, in which GH males tend to reduce the low-intensity lunging and increase the high-intensity tussling, which enhances their territorial and mating competition. We further showed that the two forms of fighting are mediated by distinct sensory and central neurons. The complex implications of these findings are discussed below.

Revised version: Regulation of physiology and behavior by social experience is an important feature across animal species. Here, we found that social experience shapes the fighting strategies of male flies, in which GH males tend to reduce the low-intensity lunging and increase the high-intensity tussling. We further showed that the two forms of fighting are mediated by distinct sensory and central neurons. The complex implications of these findings are discussed below.

Earlier version: Previous studies have found that social isolation generally enhances aggression but decreases mating competition in animal models (Palavicino-Maggio and Sengupta, 2022; Sethi et al., 2019; Yadav et al., 2024), presenting a paradox between social experience, aggression and reproductive success. Our results resolved this paradox by dissecting aggressive behaviors into the well-studied, frequently observed lunging behavior and the less frequent but high-intensity tussling behavior. We improved the behavioral paradigm for studying tussling behavior and revealed that social experience oppositely regulates lunging and tussling. It is important to note that the frequency of lunging decreases while the frequency of tussling increases from 7-day-old to 14-day-old males, suggesting a shift in fighting strategy with age. Our results further suggest that social experience may enhance male mating competition by shifting their fighting strategy to the low-frequency, high-intensity tussling, which may be more effective in combat, although direct evidence is still lacking. These results are consistent with a previous study showing a correlation between the establishment of male dominance and tussling behavior, but not lunging (Simon and Heberlein, 2020). Interestingly, we found that the mating advantage gained through social enrichment can even offset the mating disadvantage associated with aging, further supporting the vital role of shifting fighting strategies in experienced, aged males. Note that a previous study found a mating advantage in 7-day-old males compared to 2-day-old males (Lin et al., 2016), which may be due to different levels of sexual maturity, but not necessatily reflect the aging effect. These results separate the function of the two fighting forms and rectify our understanding of how social experiences regulate aggression and reproductive success.

Revised version: Previous studies have found that social isolation generally enhances aggression but decreases mating competition in animal models (Palavicino-Maggio and Sengupta, 2022; Sethi et al., 2019; Yadav et al., 2024), presenting a paradox between social experience, aggression and reproductive success. Our results provide an explanation for this paradox by dissecting aggressive behaviors into the well-studied, frequently observed lunging behavior and the less frequent but high-intensity tussling behavior. We improved the behavioral paradigm for studying tussling behavior and revealed that social experience oppositely regulates lunging and tussling. It is important to note that the frequency of lunging decreases while the frequency of tussling increases from 7-day-old to 14-day-old males, suggesting a shift in fighting strategy with age. Our results further suggest that social experience may enhance male mating competition by shifting their fighting strategy to the low-frequency, high-intensity tussling, which may be more effective in combat, although direct evidence is still lacking. These results are consistent with a previous study showing a correlation between the establishment of male dominance and tussling behavior, but not lunging (Simon and Heberlein, 2020). Interestingly, we found that the mating advantage gained through social enrichment can even offset the mating disadvantage associated with aging. Note that a previous study found a mating advantage in 7-day-old males compared to 2-day-old males (Lin et al., 2016), which may be due to different levels of sexual maturity, but not necessatily reflect the aging effect.

Earlier version: Our results revealed distinct neural circuits underlying the two forms of fighting, spanning from primary sensory neurons to central regulatory neurons. Specifically, silencing the cVA-sensing Or67d neurons reduced male lunging but not tussling, whereas inactivating Or47b neurons impaired male tussling but not lunging. Notably, Or47b neurons respond to pheromones common to both sexes, and the sensitivity of these sensory neurons is enhanced by social enrichment (Lin et al., 2016; Peng et al., 2021; Sethi et al., 2019). Unlike the Or67d-dependent lunging, which is primarily male-induced and socially inhibited, the Or47b-dependent tussling is promoted by social enrichment with either male or female flies. Recently, it has been found that female flies increase aggression towards mating pairs or mated females, depending on Or47b and Or67d, respectively (Gaspar et al., 2022; Wan et al., 2023). These results suggest a common function of Or47b and Or67d neurons in both sexes for regulating aggression. However, whether these neurons mediate different fighting patterns in females and how they respond to social experience in females remain unknown. Nevertheless, our results on these two sensory pathways in males provide crucial evidence for dissecting their functions.

Revised version: Our results revealed distinct neural circuits underlying the two forms of fighting, spanning from primary sensory neurons to central regulatory neurons. Specifically, silencing the cVA-sensing Or67d neurons reduced male lunging but not tussling, whereas inactivating Or47b neurons impaired male tussling but not lunging. Notably, Or47b neurons respond to pheromones common to both sexes, and the sensitivity of these sensory neurons is enhanced by social enrichment (Lin et al., 2016; Peng et al., 2021; Sethi et al., 2019). Unlike the Or67d-dependent lunging, which is primarily male-induced and socially inhibited, the Or47b-dependent tussling is promoted by social enrichment with either male or female flies. Recently, it has been found that female flies increase aggression towards mating pairs or mated females, depending on Or47b and Or67d, respectively (Gaspar et al., 2022; Wan et al., 2023). These results suggest a common function of Or47b and Or67d neurons in both sexes for regulating aggression. However, whether these neurons mediate different fighting patterns in females and how they respond to social experience in females remain unknown. According to the FlyWire connectome database (Dorkenwald et al., 2024; Schlegel et al., 2024), connecting Or47b and Or67d ORNs to pC1 neurons in females requires at least two intermediate neurons, suggesting a high degree of parallele processings from ORNs to pC1s.

Earlier version: Regarding central regulatory neurons, we identified a novel class of pC1 neurons, pC1SS2, that specifically promote male tussling. In male flies, there are ∼60 pairs of dsx-expressing pC1 neurons, including the well-characterized P1a neurons that also express fruM (Asahina, 2018; Jiang and Pan, 2022); however, there are only three pairs of pC1SS2neurons that specifically express dsx but not fruM, and are crucial for the high-intensity tussling in males. The three pairs of pC1SS2 neurons may have distinct function from P1a neurons in several ways. (1) P1a neurons promotes both male courtship and aggression (Hoopfer et al., 2015; Zhao et al., 2024), whereas pC1SS2 neurons specifically promote aggression but not courtship. A recent study also identified pC1SS2 neurons as aggression-promoting, but their roles in different fighting modes are not depicted (Hindmarsh Sten et al., 2025). That tussling-promoting pC1SS2 neurons specifically express dsx but not fruM is also consistent with a previous finding that dsx-positive pC1 neurons, rather that fruM-positive pC1 neurons, are crucial for male aggression (Koganezawa et al., 2016). (2) P1a neurons promotes both lunging and tussling, although their function on tussling requires additional visual stimulation. In contrast, pC1SS2 neurons specifically promote tussling, and the presence of light enhances but is not absolutely necessary for tussling. These results align with previous findings that visual stimuli amplify P1 activation and related behaviors (Hindmarsh Sten et al., 2021; Kohatsu and Yamamoto, 2015; Pan et al., 2012), and visual stimulation is important for male aggression (Duistermars et al., 2018; Ramin et al., 2014; Schretter et al., 2020). (3) Male lunging is induced after P1a activation; whereas male tussling is induced during pC1SS2 activation. These results suggest that P1a neurons may not directly regulate aggression but instead induce a persistent state that promote aggression in general, while pC1SS2 neurons may play a more direct role in mediating tussling. Future studies are needed to further dissect the seemingly parallel sensory and central neuronal pathways involved in lunging and tussling, and to explore how these pathways interact during social experience and aging to optimize reproductive success.

Revised version: Regarding central regulatory neurons, we identified a novel class of pC1 neurons, pC1SS2, that specifically promote male tussling. In male flies, there are ∼60 pairs of dsx-expressing pC1 neurons, including the well-characterized P1a neurons that also express fruM (Asahina, 2018; Jiang and Pan, 2022); however, there are only three pairs of pC1SS2 neurons that specifically express dsx but not fruM, and are crucial for the high-intensity tussling in males. The three pairs of pC1SS2 neurons may have distinct function from P1a neurons in several ways. (1) P1a neurons promotes both male courtship and aggression (Hoopfer et al., 2015; Zhao et al., 2024), whereas pC1SS2 neurons specifically promote aggression but not courtship. A recent study also identified pC1SS2 neurons as aggression-promoting, but their roles in different fighting modes are not depicted (Hindmarsh Sten et al., 2025). That tussling-promoting pC1SS2 neurons specifically express dsx but not fruM is also consistent with a previous finding that dsx-positive pC1 neurons, rather than fruM-positive pC1 neurons, are crucial for male aggression (Koganezawa et al., 2016). Moreover, the anatomical features of pC1SS2 neurons are highly similar to the male-specific aggression-promoting (MAP) neurons identified previously (Chiu et al., 2021). (2) P1a neurons promotes both lunging and tussling, although their function on tussling requires additional visual stimulation. In contrast, pC1SS2 neurons specifically promote tussling, and the presence of light enhances but is not absolutely necessary for tussling. These results align with previous findings that visual stimuli amplify P1 activation and related behaviors (Hindmarsh Sten et al., 2021; Kohatsu and Yamamoto, 2015; Pan et al., 2012), and visual stimulation is important for male aggression (Duistermars et al., 2018; Ramin et al., 2014; Schretter et al., 2020). (3) Male lunging is induced after P1a activation; whereas male tussling is induced during pC1SS2 activation. These results suggest that P1a neurons may not directly regulate aggression but instead induce a persistent state that promote aggression in general, while pC1SS2 neurons may play a more direct role in mediating tussling. Future studies are needed to further dissect the seemingly parallel sensory and central neuronal pathways involved in lunging and tussling, and to explore how these pathways interact during social experience and aging to optimize reproductive success.

Added in the revised version: Limitations of the Study

Added in the revised version: We manually analyzed aggressive behaviors and strived to ensure objectivity in behavioral analysis; however, precisely tracking and quantifying specific occurrences of lunging and tussling events for each fly over extended time periods remained methodologically challenging. Furthermore, tussling events often exhibit prolonged durations, which may also include other forms of aggression such as holding. While our results indicate an association between tussling behavior and reproductive competition, a direct link between them is lacking. Future investigations are needed to establish a causal relationship between fighting strategies and reproductive success.

Materials and Methods

Fly Stocks

Flies were maintained at 22 or 25 °C in a 12 h:12 h light: dark cycle. Canton-S flies were used as the wild-type strain and obtained from Janelia Research Campus (Pan and Baker, 2014). Or67d-GAL4 (BDSC_9998), Or65a-GAL4 (BDSC_9993), Or47b- GAL4 (BDSC_9984), IR84a-GAL4 (BDSC_41734), TK-GAL4 (BDSC_51795), R19B03-GAL4 (BDSC_49830), UAS-dTrpA1 (BDSC_26263), UAS-Kir2.1 (BDSC_6595), UAS-CsChrimson (BDSC_82181), UAS-myrGFP (BDSC_32198), retro-Tango (BDSC_99661) and trans-Tango (BDSC_77124) flies were obtained from Bloomington Drosophila Stock Center (BDSC). UAS-NaChBac (Ren et al., 2001), fruGAL4 (Stockinger et al., 2005), P1a-spGAL4 (Zhang et al., 2018), pC1SS1- spGAL4 (Wang et al., 2021), pC1SS2-spGAL4 (Wang et al., 2021), UAS-dsxM-RNAi (Peng et al., 2022) and UAS-fruMi (Meissner et al., 2016) were used as previously described. UAS-Or47b RNAi (THU_2599) was obtained from TsingHua Fly Center. Detailed genotypes of flies were listed in Table S1.

Male aggression assay

For the aggression assays presented in Figure 1, the prepared food was evenly spread into the center of behavioral chambers. There was no female in the traditional paradigm (Figure 1A, left) (Wu et al., 2020; Zhou et al., 2008), while a 3-5 day old virgin female was gently fixed in the center of the food with its abdomen outside in the modified paradigm (Figure 1A, right). We used a fixed female to restrict its position in the center of food. Then two male testers were placed into the upper chambers separated from the food (with or without a fixed female) by a film barrier. All testers were loaded by cold anesthesia. After a 30-minute adaptation, the film was gently removed to allow the two males to fall into the behavioral chamber, and the aggressive behavior was recorded for 2 hours. The aggression assays in Figure 1-figure supplement 1 was performed by using the paradigm without female, while other data were acquired using the modified paradigm with a fixed female unless otherwise stated. In some experiments, both single-housed and group-housed (30-40 males per vial) males were assayed for comparison.

We manually analyzed lunging and tussling behaviors in this study. Lunging is characterized by a male raising its forelegs and quickly striking the opponent, and each lunge typically lasts less than 0.2 seconds through detailed slow-motion analysis. Tussling is characterized by both males using their forelegs and bodies to tumble over each other, and this behavior may last from seconds to minutes. Tussling is often mixed with boxing, in which both flies rear up and strike the opponent with forelegs. Since boxing is often transient and difficult to distinguish from tussling, we referred to the mixed boxing and tussling behavior simply as tussling. As we manually analyze tussling for 2 hours for each pair of males, it is possible that we may miss some tussling events, especially those quick ones.

Earlier version: For the quantification of lunge latency and frequency in Figures 2E, 2F, 3C, 3D, and Figure 1-figures supplement 1A, 1B, we manually analyzed the videos and defined the lunge latency as the time point at which the first lunging event occurred between the two male flies after the start of the recording. We then counted the number of lunge events occurring within 10 minutes after the first lunge. For the quantification of tussle latency and frequency in Figures 1B, 1C, 1D, 2C, 2D, 3A, 3B, 3E, 3F, 3G, 3H, 4A and 4D, we manually analyzed the videos and defined the tussle latency as the time point at which the first tussling event occurred between the two male flies after the start of the recording. We defined the tussle frequency (No. of tussles in 2 hr) as the total number of tussling events that occurred within two hours from the beginning of the video.

Revised version: For the quantification of lunge latency and frequency in Figures 2E, 2F, 3C, 3D, and Figure 1-figures supplement 1A, 1B, we manually analyzed the videos and defined the lunge latency as the time point at which the first lunging event occurred between the two male flies after the start of the recording. We then counted the number of lunge events occurring within 10 minutes after the first lunge. For the quantification of tussle latency and frequency in Figures 1B, 1C, 1D, 2C, 2D, 3A, 3B, 3E, 3F, 3G, 3H, 4A and 4B, we manually analyzed the videos and defined the tussle latency as the time point at which the first tussling event occurred between the two male flies after the start of the recording. We defined the tussle frequency (No. of tussles in 2 hr) as the total number of tussling events that occurred within two hours from the beginning of the video.

The food recipes for the male aggression assay were listed below. A mixture of 2.5% yeast, 2.5% sucrose and 1% agar dissolved in apple juice (Huiyuan Apple juice concentrate) was used for Figure 1B; a mixture of 2.5% yeast, different concentrations of sucrose (0, 2.5%, 5% or 10%) and 1% agar dissolved in apple juice was used for Figure 1C; a mixture of different concentrations of yeast (0, 2.5%, 5% or 10%), 2.5% sucrose and 1% agar dissolved in apple juice was used for Figure 1D; A mixture of 10% yeast and 1% agar dissolved in apple juice was used for other experiments.

dTrpA1-mediated activation experiments

For thermogenetic activation assay showed in Figure 4-figure supplement 1, two 5-7 days old males were loaded into the circular chamber (d=15mm) under cold anesthesia and separated with a transparent film. Before the recording start, the chamber was transferred to the recording platform at a specific temperature (22°C or 30°C) for 30 minutes adaptation. After that, the film was removed, and the behavior was recorded for 1 hour. The 10-minute video after film extraction was used for behavioral analysis.

CsChrimson-mediated optogenetic activation experiments

For optogenetic activation assay showed in Figure 4 and 5, males were raised at 25°C on standard fly food, collected after eclosion and housed in the dark for 5-7 days on food containing 0.4 mM all trans-Retinal (MFCD00001550, Sigma-Aldrich, St. Louis, MO). The specific light stimulation protocols are showed in the schematic (Figure 4F). In brief, Two 5-7 days old males were loaded into the circular chamber (d=10mm) under freeze anesthesia and separated with a transparent film in the dark. Then the chamber was moved on the plane light source (IR channel: 865nm, red-light channel: 630nm) for 15 mins adaptation. Before video recording, the IR light channel was switched on and adjusted to the appropriate intensty for recording. The film was removed before the recording start, with the red-light channel switched on (constant red light, 0.02 mW/mm2) during the 2nd and 5th minutes of the experiment. The statistics for the tussling, lunging, and unilateral wing extension (UWE) are as follows. For tussling time and UWE time, we used LifesongY software (website: https://sourceforge.net/projects/lifesongy/), which is written by Dr. Chao Guo based on LifesongX (Bernstein et al., 1992), to calculate the percentage of the total time these behaviors occurred per minute. Specifically, the "tussling time (%)" refers to the percentage of time during the observation period in which tussling behavior occurred, relative to the total observation time. Similarly, "UWE time (%)" represents the percentage of time during the observation period in which UWE occurred. For lunges, we counted the number of lunging occurrences per minute.

Territorial competition experiment

For the territorial competition experiment in Figures 6A-6C, Figure 6-figure supplement 3A, the behavioral setup was produced using a 3D printer (Bambu Lab A1), as illustrated in the schematic, with a transparent cover plate placed on the top. The preparation of food and female were the same as the above aggression paradigm. The female was fixed in the center of the food and two male testers were placed separately in the two small chambers on both sides. Before initiating the experiment, the two barriers were closed. After a 30-minute adaptation, the barriers were opened to allow the males access to the central large chamber, and were closed once the males went inside, and the territorial behavior was recorded for 2 hours. The first 10 encounters occurring after 1 hour of recording were used for statistical analysis. The winning event was defined as a male driving the other male out of the food area. To discriminate between tester males, about half of them were marked on one of their wings using a red maker pen 24 hours before the experiment. For expreiments using males with different ages and/or housing conditions, we calculated the winning index as (num. of wins by males [upper panel] – num. of wins by males [lower panel])/10 encounters * 100%. For control expreiments using males of the same age and housing condition, we calculated the winning index as (num. of wins by unmarked males – num. of wins by marked males)/10 encounters * 100%, which is roughly zero due to the randomness of marking.

Mating competition experiment

For the mating competition experiment shown in Figure 6D-6H, Figure 6-figure supplement 3B, the wings of males were marked 24 hours prior to the experiment the same as described above. During the experiment, a 3-5 days old virgin female was placed into the chamber under freeze anesthesia and covered with a transparent film. Then the unmarked males were loaded into the chamber in the same manner and also covered with a transparent film. Finally, the marked males were loaded. After 30 minutes adaptation, two films were removed to allow the three flies to interact, and the behavior was recorded for 1 hour. The chambers without successful copulation during the 1-hour period were excluded from further analysis. For each condition, we used 120 efficient chambers with 10 random chambers collectively as one sample. For expreiments using males with different ages and/or housing conditions, the copulation advance index was defined as (num. of male winners [upper panel] – num. of male winners [lower panel])/10 chambers * 100%. For control expreiments using males of the same age and housing condition, the copulation advance index was calculated as (num. of winners by unmarked males – num. of winners by marked males)/10 chambers * 100%, which is roughly zero due to the randomness of marking. The winner was defined as the male that first copulated with the female.

Quantitative Real-time PCR

Earlier version: Adult fly samples were frozen in liquid nitrogen and decapitated by vortex. The heads were then separated from the bodies using metal sieves. Each sample, consisting of 40 frozen heads, was used for total RNA preparation using TRIzol reagent (15596026, Thermo Fisher Scientific, Waltham, MA). We purified total RNA using a DNA-freeTM Kit (AM1906, Thermo Fisher Scientific, Waltham, MA) and performed reverse transcription using SuperScriptTM IV (18091050, Thermo Fisher Scientific, Waltham, MA) to obtain cDNA used for templates. Quantitative PCR was performed on the Roche LightCycler® 96 Real-Time PCR machine using AceQ qPCR SYBR Green Master Mix (Q121-02, Vazyme, Nanjing). Transcript levels were analyzed by the 2^- ΔCT method using actin as control. Each sample was run in triplicate. Each experiment was repeated three times using independent sets of genetic crosses. Primers used for RT-qPCR quantification were:

Revised version: Adult fly samples were frozen in liquid nitrogen and decapitated by vortex. The heads were then separated from the bodies using metal sieves. Each sample, consisting of 40 frozen heads, was used for total RNA preparation using TRIzol reagent (15596026, Thermo Fisher Scientific, Waltham, MA). We purified total RNA using a DNA-freeTM Kit (AM1906, Thermo Fisher Scientific, Waltham, MA) and performed reverse transcription using SuperScriptTM IV (18091050, Thermo Fisher Scientific, Waltham, MA) to obtain cDNA used for templates. Quantitative PCR was performed on the Roche LightCycler® 96 Real-Time PCR machine using AceQ qPCR SYBR Green Master Mix (Q121-02, Vazyme, Nanjing). Transcript levels were analyzed by the 2^-ΔCT method using actin as control. Each sample was run in triplicate. Each experiment was repeated three times using independent sets of genetic crosses. Primers used for RT-qPCR quantification were: actin forward: 5′-GTCGCGATTTAACCGACTACCTGA-3′ actin reverse: 5′-TCTTGCTT CGAGATCCACATCTGC-3′ fru-P1 forward: 5′-GTGTGCGTACGTTTGAGTGT-3′ fru-P1 reverse: 5′-TAATCCTGTGACGTCGCCAT-3′ Or47b forward: 5′-CAAATCTCAGCCTTCTGCGG-3′ Or47b reverse: 5′-GATACTGGCACAGCAAACTCA-3′

Removed in the revised version: actin forward: 5′-GTCGCGATTTAACCGACTACCTGA-3′

Removed in the revised version: actin reverse: 5′-TCTTGCTT CGAGATCCACATCTGC-3′

Removed in the revised version: fru-P1 forward: 5′-GTGTGCGTACGTTTGAGTGT-3′

Removed in the revised version: fru-P1 reverse: 5′-TAATCCTGTGACGTCGCCAT-3′

Removed in the revised version: Or47b forward: 5′-CAAATCTCAGCCTTCTGCGG-3′

Removed in the revised version: Or47b reverse: 5′-GATACTGGCACAGCAAACTCA-3′

Tissue Dissection, Staining, and Imaging

Earlier version: Brains were dissected in Schneider’s insect medium (S2) and fixed in 4% paraformaldehyde in 0.5% Triton X-100 and 0.5% bovine serum albumin in phosphate-buffered saline (PAT) for 30 min at room temperature. After 4 × 10-min washes, tissues were blocked in 3% normal goat serum (NGS) for 60 min, then incubated in primary antibodies diluted in 3% NGS for 4 h at room temperature and 1 day at 4°C, then washed in PAT, and incubated in secondary antibodies diluted in 3% NGS for 4 h at room temperature and 1 day at 4°C. The tissues were then washed thoroughly in PAT and mounted for imaging. Primary antibodies used in this study include mouse anti-HA (BioLegend, 901501, 1:200), rabbit anti-GFP (Invitrogen, A11122, 1:1000), monoclonal mouse anti-GFP (Sigma-Aldrich, G6539, 1:500), mouse anti-Bruchpilot (Developmental Studies Hybridoma Bank, nc82, 1:50), rabbit anti-DsxM (1:500) and mouse anti-FruM (1:2000). Secondary antibodies used include donkey anti-mouse IgG conjugated to Alexa 488 (Invitrogen, A21206, 1:500), donkey anti-mouse IgG conjugated to Alexa 555 (Invitrogen, A-31570, 1:500), donkey anti-rabbit IgG conjugated to Alexa 488 (Invitrogen, A-21206, 1:500) and donkey anti-rabbit IgG conjugated to Alexa 555 (Invitrogen, A-31572, 1:500). Samples were imaged on Zeiss 900 confocal microscopes using ZEN and processed with Fiji software.

Revised version: Brains were dissected in Schneider’s insect medium (S2) and fixed in 4% paraformaldehyde in 0.5% Triton X-100 and 0.5% bovine serum albumin in phosphate-buffered saline (PAT) for 30Lmin at room temperature. After 4L×L10-min washes, tissues were blocked in 3% normal goat serum (NGS) for 60Lmin, then incubated in primary antibodies diluted in 3% NGS for 4Lh at room temperature and 1 day at 4°C, then washed in PAT, and incubated in secondary antibodies diluted in 3% NGS for 4Lh at room temperature and 1 day at 4°C. The tissues were then washed thoroughly in PAT and mounted for imaging. Primary antibodies used in this study include mouse anti-HA (BioLegend, 901501, 1:200), rabbit anti-GFP (Invitrogen, A11122, 1:1000), monoclonal mouse anti-GFP (Sigma-Aldrich, G6539, 1:500), mouse anti-Bruchpilot (Developmental Studies Hybridoma Bank, nc82, 1:50), rabbit anti-DsxM (1:500) and mouse anti-FruM (1:2000). Secondary antibodies used include donkey anti-mouse IgG conjugated to Alexa 488 (Invitrogen, A21206, 1:500), donkey anti-mouse IgG conjugated to Alexa 555 (Invitrogen, A-31570, 1:500), donkey anti-rabbit IgG conjugated to Alexa 488 (Invitrogen, A-21206, 1:500) and donkey anti-rabbit IgG conjugated to Alexa 555 (Invitrogen, A-31572, 1:500). Samples were imaged on Zeiss 900 confocal microscopes using ZEN and processed with Fiji software.

Trans-synaptic labeling by trans-Tango and retro-Tango

The trans-Tango or retro-Tango ligand with GFP was expressed under the GAL4 driver. Trans-synaptic signals were marked by a membrane-associated form of tdTomato (QUAS-mtdTomato-3xHA). The anti-HA antibody (BioLegend, 901501, 1:200) was applied to characterize the postsynaptic or presynaptic signal. For Figure 3-figure supplement 2, male testers were kept at 18 °C until eclosion and subsequently single- or group-housed at the same temperature for two weeks before dissection and imaging. For Figure 4-figure supplement 2, flies were kept at 18 °C for trans-Tango and 25 °C for retro-Tango experiments.

Brain image registration

The protocol for standard brain registration is described previously (Zhou et al., 2014). Confocal images for brains containing P1a neurons and pC1SS2 neurons were registered onto a standard reference brain (Zhou et al., 2014), using a Fiji graphical user interface as described previously (Ostrovsky et al., 2013).

Statistics

Experimental flies and genetic controls were tested under the same conditions, and data were collected from at least two independent experiments. Statistical analysis was performed using GraphPad Prism 9 (https://www.graphpad.com/scientific-software/prism/) as indicated in each figure legend. Data were first verified for normal distribution by the D’Agostino–Pearson normality test. For normally-distributed data, Student’s t test was used for pairwise comparisons, and one-way ANOVA was used for comparisons among multiple groups, followed by Tukey’s multiple comparisons. For data not normally-distributed, the Mann-Whitney test was used for pairwise comparisons, and the Kruskal–Wallis test was used for comparison among multiple groups, followed by Dunn’s multiple comparisons. For group matched data that were normally-distributed, the Paired t test signed rank test was used for pairwise comparisons, and RM one-way ANOVA was used for comparison among multiple groups, followed by RM one-way ANOVA multiple comparisons. For group matched data that were not normally-distributed, the Wilcoxon matched-pairs signed rank test was used for pairwise comparisons, and the Friedman test was used for comparison among multiple groups, followed by Dunn’s multiple comparisons. For qPCR experiments, the average relative expression of three independent experiments was analyzed using Mann-Whitney test. For one group compared with zero, we used one sample t test. The Chi-square test was used to compare male tussling percentage between two groups.